# Welcome to your homework Workshop Attendance!

In this homework we will go through some exercise utilizing almost all of the material that have been taught in the workshop. This is in no way a really thorough test, you can excercise more using programming test website such as hackerrank and leetcode. The difficulties that is written in each problem is subjective, I rate them based on a beginner level.

Rules:
- Don't change the code that has been written before
- Match the variable for the return output to prevent things not working
- There are hidden test that are not written in this notebook, so make sure your function work with different inputs
- Don't copy other people answer, take your time, use this moment to study more]
- If you find an answer from stackoverflow, make sure you don't blatantly copy paste it, understand and change some things out

Tips:
- Try to understand the problem before coding right in
- Read the module once again
- Before asking a friend, try googling things out
- Never stop learning
- Have fun !

```python
# Initial Import
import numpy as np
```

## Problem 1 : Count A
Difficulty: 1/10 <br>
Point: 10

In this problem, you need to make function that count the amount of "A" or capital "a" inside a string and returning the amount of "A"

Example:
count_A("ABCA") return 2, because there are 2 "A" inside the string

Hint:
- How do you iterate through a list?
- How do you implement logic?

<!-- grade_id: cell-fd29794207e80b4c, points: 0, grade: false, solution: true -->
```python
def count_A(s):
    '''
    Count the amount of "A" inside of input s
    '''
    ### DON'T CHANGE THE CODE PROVIDED
    count = 0
    
    # YOUR CODE HERE
    
    return count
```

<!-- grade_id: cell-e600a2f61f62fa5f, points: 10, grade: true, solution: false -->
```python
assert count_A("ABC") == 1
assert count_A("AaAaA") == 3
assert count_A("") == 0
```

## Problem 2 : Flatten List of List
Difficulty: 5/10<br>
Point: 30

In this problem, you need to make function that flatten a list of list, **becareful for input with multiple list of list**

Example:
flatten_list([1,[2,3],[4,5],6]) return [1,2,3,4,5,6] because it remove the list of list and return the whole thing inside the list

Hint:
- You can iterate over an iteration
- Use [isinstance()](https://www.w3schools.com/python/ref_func_isinstance.asp)

<!-- grade_id: cell-6779c45fd2b1e0a4, points: 0, grade: false, solution: true -->
```python
def flatten_list(l):
    ### DON'T CHANGE THE CODE PROVIDED
    output = []
    
    # YOUR CODE HERE
    
    return output
```

<!-- grade_id: cell-ae8e23062c41cdfb, points: 30, grade: true, solution: false -->
```python
assert flatten_list([1,[2,3]]) == [1,2,3]
assert flatten_list([1,[2,3],[]]) == [1,2,3]
assert flatten_list([[1],[2],[3]]) == [1,2,3]
```

## Problem 3 : Compute Force Between Two Particle
Difficulty: 3/10<br>
Point: 10

In this problem, make a function that satisfy the formula of force between two charged particle, lets test your physics a little bit<br>
$$F = \frac{k.Q1.Q2}{r^2} $$

Example:
compute_force(1*10e-9,2,3) return 200

Hint:
- Becareful with how you write the formula in Python
- Becareful with 0

<!-- grade_id: cell-c8e698284b9b92da, points: 0, grade: false, solution: true -->
```python
def compute_force(q1,q2,r):
    ### DON'T CHANGE THE CODE PROVIDED
    k = 9*10e9

    # YOUR CODE HERE
    
    return F
```

<!-- grade_id: cell-d1b7daba6b413a5c, points: 10, grade: true, solution: false -->
```python
assert compute_force(1,2,3) == 2*10e9
assert compute_force(1*10e-9,2,3) == 200
assert compute_force(1,2,0) == float("infinity")
```

## Problem 4: Compute Area of Rectangle
Difficulty: 4/10<br>
Point: 20

In this problem you need to compute the area of a rectangle using the coordinate of only two corner that is placed diagonally (not sorted and can be either left or right diagonal).

Hint:
- How do you calculate the area? Maybe there is a formula?
- Becareful with negative output
- If only there is a way to make negative value to positive

<!-- grade_id: cell-3ae615e14966347e, points: 0, grade: false, solution: true -->
```python
def compute_area(pair_1,pair_2):
    ### DON'T CHANGE THE CODE PROVIDED
    x1,y1 = pair_1
    x2,y2 = pair_2
    
    # YOUR CODE HERE
    
    return area
```

<!-- grade_id: cell-4f0bd17c7c530bb1, points: 20, grade: true, solution: false -->
```python
assert compute_area([1,2],[3,4]) == 4
assert compute_area([3,4],[1,2]) == 4
assert compute_area([2,2],[6,-2]) == 16
assert compute_area([6,-2],[2,2]) == 16
```

## Problem 5: Array Normalization
Difficulty: 7/10<br>
Point: 30

In this problem you need to make a function that normalize an array to a scale of 0-1. Process of normalization is usually used when your data value is too high or low, this is useful for data analysis and machine learning. The formula for normalization can be written as

$$x_{Norm} = \sum_{n=0}^{N} \frac{x_n - x_{min}}{x_{max}-x_{min}} $$

Hint:
- Use numpy for this problem, it is already imported on the top
- Becareful with 0

<!-- grade_id: cell-8fd6cbadf84b22df, points: 0, grade: false, solution: true -->
```python
def normalize(x):
    
    # YOUR CODE HERE
    
    return x_norm
```

<!-- grade_id: cell-218523a772943f6f, points: 30, grade: true, solution: false -->
```python
np.testing.assert_allclose(normalize([]), [])
np.testing.assert_allclose(normalize([0]), [0])
np.testing.assert_allclose(normalize([0, 0]), [0, 0])
np.testing.assert_allclose(normalize([1]), [1])
np.testing.assert_allclose(normalize([1, 1]), [1, 1])
np.testing.assert_allclose(
        normalize(np.arange(10)),
        [
            0,
            0.11111111,
            0.22222222,
            0.33333333,
            0.44444444,
            0.55555556,
            0.66666667,
            0.77777778,
            0.88888889,
            1,
        ],
    )

np.testing.assert_allclose(
    normalize(np.arange(10, 0, -1)),
    [
        1.0,
        0.88888889,
        0.77777778,
        0.66666667,
        0.55555556,
        0.44444444,
        0.33333333,
        0.22222222,
        0.11111111,
        0.0,
    ],
)

np.testing.assert_allclose(
    normalize(np.arange(0, -10, -1)),
    [
        1.0,
        0.88888889,
        0.77777778,
        0.66666667,
        0.55555556,
        0.44444444,
        0.33333333,
        0.22222222,
        0.11111111,
        0.0,
    ],
)
```

# Thats It!
Now run everything and make sure everything is working well, **MAKE SURE YOU SAVE THE FILE** before downloading the file, to download you need to go to "File" -> "Download as" -> ".ipynb"/"notebook", it maybe different in your IDE but in the end you need to save it in notebook too 

**Save it using this naming format: YOUR_FULL_NAME.ipynb**